# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = ""
COLLABORATORS = ""
```

---

## Zestaw zadań 1: Elementy statystyki opisowej 1 (Notatnik 3/4)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import dataframe_image as dfi
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Wczytanie danych
Wczytaj dane z przygotowanego wcześniej pliku dane.csv do tablicy `eruptions_data` w taki sposób, aby kolumna Eruption Number stanowiła indeks tablicy.

Wyświetl przygotowaną tablicę.

<!-- grade_id: cell-184ec13f50fb3bfe, points: 0, grade: true, solution: true -->
```python
eruptions_data = pd.read_csv("eruptions_data.csv", index_col="Eruption Number")

eruptions_data.head()
```

Output:
```
Volcano Name Eruption Category Evidence Type  \
Eruption Number                                                
22486               Cotopaxi         Confirmed  Observations   
22481                   Taal         Confirmed  Observations   
22458              Turrialba         Confirmed  Observations   
22453                 Ulawun         Confirmed  Observations   
22454                  Raung         Confirmed  Observations   

                      Evidence Method  VEI  Start Year  Start Month  \
Eruption Number                                                       
22486                        Reported    2        2022         10.0   
22481                        Reported    1        2022         10.0   
22458                        Reported    1        2022          7.0   
22453                        Reported    2        2022          6.0   
22454            Satellite (infrared)    2        2022          5.0   

                 Start Day  End Year  End Month  End Day  
Eruption N
```

### Zadanie 4: Konstrukcja szeregów rodzielczych - bardziej skomplikowane tablice [10 pkt]

#### a) Forma tabelaryczna
Przygotuj szeregi rozdzielcze VEI w formie tablicy `vei_data` zawierającej następujące elementy:

Indeks:
- VEI - wartości VEI posortowane od najmniejszej do największej.

Kolumny: </br>
- Nagłówki 1 poziomu:
    - Eruption category - All, Confirmed Eruption, Uncertain Eruption
- Nagłówki 2 poziomu (dla każdego z nagłówków 1 poziomu ten sam zestaw kolumn):
    - f - zliczenia poszczególnych kategorii (ang. frequency, absolute frequency),
    - cf - zliczenia skumulowane poszczególnych kategorii (ang. cumulative frequency, cumulative absolute frequency),
    - rf - częstość zliczeń poszczególnych kategorii (ang. relative frequency),
    - crf - skumulowana częstość zliczeń poszczególnych kategorii (ang. cumulative relative frequency).

Wyświetl przygotowane tablicę w taki sposób, aby liczby w kolumnach f i cf oraz indeks tablicy VEI wyświetlały się bez miejsc po przecinku, a liczby z kolumn rf i crf z dokładnością do 4 miejsca po przecinku.

<!-- grade_id: cell-10bd4bdf7a80603c, points: 0, grade: false, solution: true -->
```python
header = pd.MultiIndex.from_product([['All eruptions','Confirmed Eruptions', 'Uncertain Eruptions'],
                                     ['f', 'cf', 'rf', 'crf']],
                                    names=['loc','S']) # do każdej naglowka gornego dodaje wszystkie z nahlowka dolnego

data = eruptions_data.copy()
data_confirmed = data[data["Eruption Category"] == "Confirmed"]
data_confirmed = data[data["Eruption Category"] == "Confirmed"]
data_unconfirmed = data[data["Eruption Category"] == "Uncertain"]

vei_data = pd.DataFrame({
    "VEI": range(0, 8),
    "f": data.groupby("VEI")["Volcano Name"].count(), # zliczanie poszczrgolnych kategorii
    "cf": np.cumsum(data.groupby("VEI")["Volcano Name"].count()), # suma skumulowana
    "rf": data.groupby("VEI")["Volcano Name"].count() / len(data["VEI"]), # częstość zliczeń poszczególnych kategorii 
    "crf": np.cumsum(data.groupby("VEI")["Volcano Name"].count()) / len(data), #skumulowana częstość zliczeń poszczególnych kategorii
    
    "f": data_confirmed.groupby("VEI")["Volcano Name"].count(), # zliczanie poszczrgolnych kategorii
    "cf": np.cumsum(data_confirmed.groupby("VEI")["Volcano Name"].count()), # suma skumulowana
    "rf": data_confirmed.groupby("VEI")["Volcano Name"].count() / len(data_confirmed), # częstość zliczeń poszczególnych kategorii 
    "crf": np.cumsum(data_confirmed.groupby("VEI")["Volcano Name"].count()) / len(data_confirmed) ,#skumulowana częstość zliczeń poszczególnych kategorii

    "f": data_unconfirmed.groupby("VEI")["Volcano Name"].count(), # zliczanie poszczrgolnych kategorii
    "cf": np.cumsum(data_unconfirmed.groupby("VEI")["Volcano Name"].count()), # suma skumulowana
    "rf": data_unconfirmed.groupby("VEI")["Volcano Name"].count() / len(data_unconfirmed), # częstość zliczeń poszczególnych kategorii 
    "crf": np.cumsum(data_unconfirmed.groupby("VEI")["Volcano Name"].count()) / len(data_unconfirmed)
    
    # uzupelnij None
},
                  index="VEI",
                  columns=header)
```

Output:
```
File c:\Users\Kacper\Desktop\sem_VI\analiza_danych_python\venv\lib\site-packages\pandas\core\indexes\base.py:716, in Index._with_infer(cls, *args, **kwargs)
    714 with warnings.catch_warnings():
    715     warnings.filterwarnings("ignore", ".*the Index constructor", FutureWarning)
--> 716     result = cls(*args, **kwargs)
    718 if result.dtype == _dtype_obj and not result._is_multi:
    719     # error: Argument 1 to "maybe_convert_objects" has incompatible type
    720     # "Union[ExtensionArray, ndarray[Any, Any]]"; expected
    721     # "ndarray[Any, Any]"
    722     values = lib.maybe_convert_objects(result._values)  # type: ignore[arg-type]

File c:\Users\Kacper\Desktop\sem_VI\analiza_danych_python\venv\lib\site-packages\pandas\core\indexes\base.py:545, in Index.__new__(cls, data, dtype, copy, name, tupleize_cols, **kwargs)
    542     return klass._simple_new(arr, name)
    544 elif is_scalar(data):
--> 545     raise cls._scalar_data_error(data)
    546 elif hasattr(data,
```

<!-- grade_id: cell-37f309620c903da9, points: 5, grade: true, solution: false -->
```python
## Komórka testowa
# Test 1: Indeks
assert (vei_data.index.name == 'VEI') == True
assert list(vei_data.index) == [0, 1, 2, 3, 4, 5, 6, 7]
# Test 2: Kolumny
assert len(set(list(vei_data.columns.levels[0])) & set(['All eruptions', 'Confirmed Eruptions', 'Uncertain Eruptions'])) == 3
assert set(list(vei_data.columns.levels[1])) == set(['f', 'cf', 'rf', 'crf'])
# Test 3: Wartości w kolumnach f
assert list(vei_data['All eruptions']['f'].values) == [1007, 1384, 3991, 1152, 511, 180, 53, 7]
assert list(vei_data['Confirmed Eruptions']['f'].values) == [803, 1227, 3727, 1133, 510, 180, 53, 7]
assert list(vei_data['Uncertain Eruptions']['f'].values[:5]) == [204, 157, 264, 19, 1] 
assert np.all(np.isnan(vei_data['Uncertain Eruptions']['f'].values[5:]))
```

#### b) Forma graficzna
Zwizualizuj dane z przygotowanych tablic za pomocą wykresów słupkowych zliczeń (lewa oś wykresu) z nałożonymi na nie łamanymi częstości skumulowanej (prawa oś wykresu). Zestaw wykresy na pojedynczej figurze (siatka wykresów 1x3). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo ustaw:
- identyczne zakresy wartości na odpowiadających sobie osiach wszystkich trzech wykresów,
- różne od domyślnych i od siebie nawzajem kolory wykresu słupkowego oraz łamanej częstości,
- kolory lewej i prawej osi oraz ich opisów identyczne z kolorem elementu wykresu, którego dotyczą,
- format wyświetlania wartości VEI bez miejsc po przecinku.

<!-- grade_id: cell-4291e1fff2828e80, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Zadanie 5: Rodzaje metod wykrywania erupcji [10 pkt]

Przygotuj tablice `evidence_data` zawiarającą informacje o tym ile erupcji o poszczególnych wartościach VEI zostało potwierdzonych przy użyciu poszczególnych metod weryfikacji.
Tablica powinna zawierać następujące elementy:

Indeks:
- Nagłówki 1 poziomu
    - kategoria dowodów
- Nagłówki 2 poziomu:
    - metoda, która dostarczyła dowodów

Kolumny: </br>
- VEI - wartości VEI posortowane od najmniejszej do największej oraz kolumna "All", w której znajdują się dane dotyczące wszystkich erupcji, bez podziału na VEI.

Wyświetl przygotowane tablicę w taki sposób, aby zawarte w tablicy liczby (w tym wartości VEI w nagłówkach kolumn) wyświetlały się bez miejsc po przecinku.

<!-- grade_id: cell-57304971296e4e92, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-49e4af55295b70a6, points: 10, grade: true, solution: false -->
```python
# Test 1 Kształt tablicy
assert evidence_data.shape == (22, 9)
# # Test 2 Kolumny
assert len(set(list(evidence_data.columns)) & set([0.0, 1.0, 2.0, 3.0, 4.0, 5.0, 6.0, 7.0, 'All'])) == 9
# Test 3 Indeks
assert (eruptions_data.index.name == 'Eruption Number') == True
assert len(set(list(evidence_data.index)) & set([('Observations', 'Reported'), ('Observations', 'Satellite (infrared)'), ('Observations', 'Satellite (visual)'),
                                ('Observations', 'Seismicity'), ('Observations', 'Hydrophonic'), ('Observations', 'Photo / Video'), ('Observations', 'Aviation'),
                                ('Uncertain', 'Unspecified'), ('Isotopic', 'Uranium-series'), ('Isotopic', '14C (calibrated)'), ('Isotopic', '14C (uncalibrated)'),
                                ('Isotopic', 'Cosmic Ray Exposure'), ('Isotopic', 'Ar/Ar'), ('Isotopic', 'K/Ar'), ('Correlation', 'Magnetism'), ('Correlation', 'Tephrochronology'),
                                ('Correlation', 'Anthropology'), ('Sidereal', 'Varve Count'), ('Sidereal', 'Dendrochronology'), ('Sidereal', 'Ice Core'), ('Radiogenic', 'Fission track'),
                                ('Radiogenic', 'Thermoluminescence')])) == 22
# Test 4 Kilka losowo wybranych wartości w tablicy
assert evidence_data.loc[('Observations', 'Reported'), 0] == 376
assert evidence_data.loc[('Uncertain', 'Unspecified'), 3] == 22
assert evidence_data.loc[('Sidereal', 'Ice Core'), 7] == 2
assert evidence_data.loc[('Isotopic', '14C (uncalibrated)'), 5] == 41
assert evidence_data.loc[('Correlation', 'Anthropology'), 6] == 0
```

### Zadanie 6: Wizualizacja rozkładu zarejestrowanych erupcji w czasie [2 pkt]

Korzystając z funkcji `seaborn.stripplot()` zwizualizuj jak rozkładały się w czasie wystąpienia zarejestrowanych erupcji. Do rozróżnienia erupcji w czasie wykorzystaj rok, w którym nastąpił ich początek.

Wyniki zestaw na pojedynczej figurze (siatka wykresów 2x1). Na górnym wykresie przedstaw dane rozróżnione tylko po kątem VEI, a na dolnym wykresie przedstaw dane rozróżnione zarówno po VEI jak i rodzaju dowodu ich wystąpienia (Evidence Type). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- dodaj legendę do dolnego wykresu w taki sposób, aby nie zasłaniała znajdujących nię na nim danych,
- ustaw format wyświetlania wartości VEI bez miejsc po przecinku.

<!-- grade_id: cell-02f74843f3b61c8d, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```